---
title: 'Understanding Your Data'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Handling Time Series Data

```yaml
type: VideoExercise
key: 6baa06f0f9
xp: 50
```

`@projector_key`
043ea510d97468d9de9662eddff29a67
